```python
# This cell is used for creating a button that hides/unhides code cells to quickly look only the results.
# Works only with Jupyter Notebooks.

from IPython.display import HTML

HTML('''<script>
code_show=true;
function code_toggle() {
if (code_show){
$('div.input').hide();
} else {
$('div.input').show();
}
code_show = !code_show
}
$( document ).ready(code_toggle);
</script>
<form action="javascript:code_toggle()"><input type="submit" value="Click here to toggle on/off the raw code."></form>''')
```

```python
# Description:
#   Exercise5 notebook.
#
# Copyright (C) 2018 Santiago Cortes, Juha Ylioinas
#
# This software is distributed under the GNU General Public 
# Licence (version 2 or later); please refer to the file 
# Licence.txt, included with the software, for details.

# Preparations
import os
import numpy as np

# Select data directory
if os.path.isdir('/coursedata'):
    # JupyterHub
    course_data_dir = '/coursedata'
elif os.path.isdir('../../../coursedata'):
    # Local installation
    course_data_dir = '../../../coursedata'
else:
    # Docker
    course_data_dir = '/home/jovyan/work/coursedata/'

print('The data directory is %s' % course_data_dir)
data_dir = os.path.join(course_data_dir, 'exercise-05-data/')
print('Data stored in %s' % data_dir)
```

# CS-E4850 Computer Vision Exercise Round 2

The exercises should be solved and the solutions submitted via Aalto JupyterHub by the deadline. 

**Deliverables:**
- **PDF** with your solutions to any pen-and-paper tasks (scanned or typeset, e.g., LaTeX). Make sure the PDF is clearly readable.
- **Jupyter notebook** (.ipynb) containing your solutions to the programming tasks. Do not change the name of the notebook file. It may result in 0 points for the exercise.

**Important:**
- Fill only the cells marked with `# YOUR CODE HERE`. Do not change function signatures.
- You may add extra cells for your own tests, but **do not** overwrite global variables or edit locked cells.
- **Never create new cells by menu commands "Edit/Copy Cells" and "Edit/Paste Cells ..."**. These commands create cells with duplicate ids and make autograding impossible. Use menu commands "Insert/Insert Cell ..." or the button with a plus sign to insert new cells.
- **All notebooks contain hidden tests** which are used for grading. They are hidden inside read-only cells. Therefore, **the read-only cells should never be removed.** 
- **Note:** Visible tests mainly check the shapes and data types of your function’s output. Hidden tests check the correctness of your solution more thoroughly. Passing the visible tests does not guarantee full points for the exercise.
- **Google Colab warning:** Uploading your assignment notebooks to Colab may cause problems. Colab can overwrite notebook metadata and break the autograding. To avoid this, we recommend copy-pasting your code into the notebooks fetched on JupyterHub. Sorry for the inconvenience.
- Running the cells in mixed order (which quite often happens while trying different things and debugging) may cause errors. While working on a particular cell be sure that you have freshly run all its preceding cells belonging to the same exercise.
- **Before submitting**, simply run all the cells of the notebook (for example, select "Restart & Run All" in the menu) and check that all the cells run properly.
- **Remember to submit your assignment!**


## Exercises 1 - 5

Remember to do the pen and paper assignments given in Exercise02.pdf.



## Exercise 6 - Project a 3D point to pixel coordinates (5 points)

### Task
Implement `project_point(P, XYZ)` function, which projects a point from world coordinates to pixel coordinates. The function takes as an input a 3×4 camera projection matrix `P` and a single 3D point
`XYZ` of shape `(3,)`, and returns its pixel coordinates `(u, v)` as a float array of shape `(2,)`.

**Input**
- `P`: `np.ndarray` with shape `(3, 4)`
- `XYZ`: `np.ndarray` with shape `(3,)`

**Output**
- `np.ndarray` with shape `(2,)` containing `(u, v)`

**Hint**
- Use homogeneous coordinates for the 3D point, and remember to give the returned $u$ and $v$ in non-homogeneous coordinates.

<!-- grade_id: cell-acbc1e3ec1cd008a, points: 0, grade: false, solution: true -->
```python
def project_point(P, XYZ):
    
    ##--your-code-starts-here--##
    # YOUR CODE HERE
    raise NotImplementedError()
    ##--your-code-ends-here--##
    
    return np.array([u, v], dtype=float)
```

<!-- grade_id: cell-203a759df5ab5640, points: 0, grade: true, solution: false -->
```python
# Visible tests

P = np.array([[800.,  0., 320.,   0.],
             [  0., 800., 240.,   0.],
             [  0.,   0.,   1.,   0.]])

uv1 = project_point(P, np.array([0.0, 0.0, 2.0]))
uv2 = project_point(P, np.array([0.5, 0.0, 2.0]))

# Basic shape/type checks
assert isinstance(uv1, np.ndarray), f"uv1 must be a numpy.ndarray, got {type(uv1)}"
assert uv1.shape == (2,), f"uv1 must have shape (2,), got {uv1.shape}"

assert isinstance(uv2, np.ndarray), f"uv2 must be a numpy.ndarray, got {type(uv2)}"
assert uv2.shape == (2,), f"uv2 must have shape (2,), got {uv2.shape}"

# Sanity check on part of the result
_expected_u2 = 520
returned_u2 = uv2[0]

assert np.allclose(returned_u2, _expected_u2, atol=1e-6), "Wrong pixel u coordinate."

print("All visible tests passed!")
```

<!-- grade_id: cell-000bef90cf660077, points: 2, grade: true, solution: false -->
```python
# HIDDEN TEST CELL
# This cell contains hidden tests for grading after the deadline.
# Please do not remove or modify this cell.

```

<!-- grade_id: cell-96b777538c49eeec, points: 3, grade: true, solution: false -->
```python
# HIDDEN TEST CELL
# This cell contains hidden tests for grading after the deadline.
# Please do not remove or modify this cell.

```

```python

```